```python
# Initialize Otter
import otter
grader = otter.Notebook("pdf_submission_test.ipynb")
```

## Question 1

Write a function called `add` that takes two integers and returns the sum.

```python
def add(a, b):
    return a + b
```

```python
grader.check("Question 1")
```

<!-- BEGIN QUESTION -->

## Question 2

Describe how the function you wrote in **Question 1** works.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

## Question 3

Describe how you feel right now

_Type your answer here, replacing this text._

<!-- END QUESTION -->



## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for Question 1 (points: 1)
```
>>> add(3, 4)
7
```
